## 1-minute introduction to Jupyter ##

A Jupyter notebook consists of cells. Each cell contains either text or code.

A text cell will not have any text to the left of the cell. A code cell has `In [ ]:` to the left of the cell.

If the cell contains code, you can edit it. Press <kbd>Enter</kbd> to edit the selected cell. While editing the code, press <kbd>Enter</kbd> to create a new line, or <kbd>Shift</kbd>+<kbd>Enter</kbd> to run the code. If you are not editing the code, select a cell and press <kbd>Ctrl</kbd>+<kbd>Enter</kbd> to run the code.

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Assignment 9: Recursion

In this assignment, you should write your code in a **readable** way, and **modularise** chunks of code that would otherwise be repeated.

Modularising a section of program code means to reorganise it in a way that allows it to be reused in another part of the code easily, usually in the form of a function.

Your function definitions should have appropriate **docstrings**.

## Part 1: Shuffling

A simple recursive function, `shuffle(inputlist)`, shuffles a list of numbers is as follows:

1. Let `inputlist` have length `n`, and its indexes running from `0` to `n - 1`.
2. `pop()` a random `list` element from `inputlist` at index `x`, where `0` <= `x` <= `n - 1`. Append this element to a new `list`, `new_list`.
3. Call `shuffle()` again with the remaining list, and append the result to `new_list`.

Implement `shuffle()` below. Include a simple validation to verify that the input is a `list` type.

You may assume that all the elements are `int`. You may use the `randint` function from the `random` module to generate a random index.

<!-- grade_id: cell-3e0505a5dd883ba4, points: 0, grade: false, solution: true -->
```python
def shuffle(inputlist):
    
```

<!-- grade_id: cell-8774eb740517b1f7, points: 1, grade: true, solution: false -->
```python
# Autograder test for valid docstring
import string

mark = '✓'
docstring = shuffle.__doc__
for char in string.whitespace:
    docstring = docstring.replace(char,'')

try:
    assert docstring != '', \
        'No docstring defined for shuffle()'
except AssertionError as e:
    mark = '✗'
    print(e)
# Hidden test!

result = f'[{mark}] Appropriate docstring  '
print(result)
```

<!-- grade_id: cell-c0f7760ca7850727, points: 1, grade: true, solution: false -->
```python
# Autograder test for output

mark = '✓'
inputlist = list(range(1,10))
_inputlist = inputlist.copy()
outputlist = shuffle(_inputlist)

try:
    assert type(outputlist) == list, \
        f'Return value should be a list, not {type(outputlist)}'
    assert len(outputlist) == len(inputlist), \
        'Returned list does not have the same number of '\
        'elements as input list'
    assert set(outputlist) == set(inputlist), \
        'Returned list does not have the same elements as input list'
    assert outputlist != inputlist, \
        'Returned list should not be the same as input list'
except AssertionError as e:
    mark = '✗'
    print(e)

try:
    assert outputlist != inputlist, \
        'Returned list should not be the same as input list'
except AssertionError as e:
    mark = '✗'
    print(e)

result = f'[{mark}] Correct output  '
print(result)
```

<!-- grade_id: cell-c9be70eedcf413a0, points: 1, grade: true, solution: false -->
```python
# Autograding test for recursive call
def func(list_):
    func.counter += 1
    return list_
func.counter = 0

mark = '✓'
__shuffle, shuffle = shuffle, func

inputlist = [1, 2, 3, 4, 5]
result = __shuffle(inputlist)
shuffle = __shuffle

try:
    assert func.counter > 0, \
        'shuffle() was not called recursively within itself'
except AssertionError as e:
    mark = '✗'
    print(e)

result = f'[{mark}] Correct self-invocation  '
print(result)
```

<!-- grade_id: cell-e4894f62c4e6dd9a, points: 1, grade: true, solution: false -->
```python
# Autograding test for recursion
import functools

mark = '✓'
__shuffle = shuffle
def counter(func, **kwargs):
    @functools.wraps(func)
    def wrapper(*args, **kwargs):
        counter.called += 1
        counter.depth += 1
        if counter.debug:
            print(f'Calling shuffle(), '
                  f'depth {counter.depth}, '
                  f'calls {counter.called}, '
                  f'args {args[0]}')
        result = func(*args, **kwargs)
        counter.depth -= 1
        return result
    counter.debug = kwargs.get('debug', False)
    return wrapper
counter.called, counter.depth = 0, 0
counter.depth = 0

shuffle = counter(shuffle, debug=False)

inputlist = list(range(1,11))
return_list = shuffle(inputlist)
try:
    assert counter.called >= 9, \
        'shuffle() should be called at least 9 times for a '\
        f'10-element list, was called {counter.called} times instead'
except AssertionError as e:
    mark = '✗'
    print(e)
finally:
    shuffle = __shuffle

result = f'[{mark}] Correct recursion depth  '
print(result)
```

<!-- grade_id: cell-05cbacc8ab2c0781, points: 1, grade: true, solution: false -->
```python
# Autograding test for base case
mark = '✓'
try:
    assert shuffle([1]) == [1], \
        'shuffle() does not seem to handle a base case '\
        '(single-element list) properly'
except AssertionError as e:
    mark = '✗'
    print(e)

result = f'[{mark}] Correct base case handling  '
print(result)
```

## Part 2: Write a palindrome checker

A palindrome is a sentence or number with a sequence of letters, such that it reads the same way forwards and backwards.

Write a function, `is_palindrome()`, that checks whether an input string is a valid palindrome. Use a recursive algorithm to do so.

1. The result should be case-insensitive
2. The function should ignore punctuation and whitespace

**Hint 1:** A palindrome should have the first and last letters be the same.

**Hint 2:** If a string is a valid palindrome, the same string without its first and last letters is also a palindrome.

<!-- grade_id: cell-be18132fed12a054, points: 0, grade: false, solution: true -->
```python
import string


def is_palindrome(inputstr):
    
    
is_palindrome('A man, a plan, a canal, Panama!')
```

<!-- grade_id: cell-54cf900b3b0f9ff5, points: 1, grade: true, solution: false -->
```python
# Autograder test for valid docstring
import string

mark = '✓'
docstring = is_palindrome.__doc__
for char in string.whitespace:
    docstring = docstring.replace(char,'')

try:
    assert docstring != '', \
        'No docstring defined for shuffle()'
except AssertionError as e:
    mark = '✗'
    print(e)
# Hidden test!

result = f'[{mark}] Appropriate docstring  '
print(result)
```

<!-- grade_id: cell-cd4a2454d599361d, points: 1, grade: true, solution: false -->
```python
# Autograder test for output

mark = '✓'
testdata = [('ABBA', True),
            ('ABBA! ', True),
            ('Able was I, ere I saw Elba', True),
            ('A man, a plan, a canal, Panama!', True),
            ('Not a palindrome', False)
            ]

for seq, ans in testdata:
    try:
        result = is_palindrome(seq)
        assert result == ans, \
            f'is_palindrome("{seq}") is {result} but should be {ans}'
    except AssertionError as e:
        mark = '✗'
        print(e)

result = f'[{mark}] Correct output  '
print(result)
```

<!-- grade_id: cell-c4e10aac5dc5ffdc, points: 1, grade: true, solution: false -->
```python
# Autograding test for recursive call
def func(list_):
    func.counter += 1
    return list_
func.counter = 0

mark = '✓'
__is_palindrome, is_palindrome = is_palindrome, func

inputstr = 'apalindrome'
result = __is_palindrome(inputstr)
is_palindrome = __is_palindrome

try:
    assert func.counter > 0, \
        'is_palindrome() was not called recursively within itself'
except AssertionError as e:
    mark = '✗'
    print(e)

result = f'[{mark}] Correct self-invocation  '
print(result)
```

<!-- grade_id: cell-affbae9f9bc0ae5e, points: 1, grade: true, solution: false -->
```python
# Autograding test for recursion
import functools

mark = '✓'
__is_palindrome = is_palindrome
def counter(func, **kwargs):
    @functools.wraps(func)
    def wrapper(*args, **kwargs):
        counter.called += 1
        counter.depth += 1
        if counter.debug:
            print(f'Calling is_palindrome(), '
                  f'depth {counter.depth}, '
                  f'calls {counter.called}, '
                  f'args {args[0]}')
        result = func(*args, **kwargs)
        counter.depth -= 1
        return result
    counter.debug = kwargs.get('debug', False)
    return wrapper
counter.called, counter.depth = 0, 0
counter.depth = 0

is_palindrome = counter(is_palindrome, debug=False)

inputstr = 'A man, a plan, a canal, Panama!'
return_str = is_palindrome(inputstr)
try:
    assert counter.called >= 6, \
        'is_palindrome() should be called at least 10 times for a '\
        f'21-element str, was called {counter.called} times instead'
except AssertionError as e:
    mark = '✗'
    print(e)
finally:
    is_palindrome = __is_palindrome

result = f'[{mark}] Correct recursion depth  '
print(result)
```

<!-- grade_id: cell-a9aaf4942f8dbd3a, points: 1, grade: true, solution: false -->
```python
# Autograding test for base case
mark = '✓'
try:
    assert is_palindrome('') == True, \
        'is_palindrome() does not seem to handle a base case '\
        '(empty string) properly'
    assert is_palindrome('a') == True, \
        'is_palindrome() does not seem to handle a base case '\
        '(single char) properly'        
except AssertionError as e:
    mark = '✗'
    print(e)

result = f'[{mark}] Correct base case handling  '
print(result)
```

# Feedback and suggestions

Any feedback or suggestions for this assignment?

<!-- grade_id: cell-63f8f1cb7ff65bb1, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE