## 1-minute introduction to Jupyter ##

A Jupyter notebook consists of cells. Each cell contains either text or code.

A text cell will not have any text to the left of the cell. A code cell has `In [ ]:` to the left of the cell.

If the cell contains code, you can edit it. Press <kbd>Enter</kbd> to edit the selected cell. While editing the code, press <kbd>Enter</kbd> to create a new line, or <kbd>Shift</kbd>+<kbd>Enter</kbd> to run the code. If you are not editing the code, select a cell and press <kbd>Ctrl</kbd>+<kbd>Enter</kbd> to run the code.

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Assignment 11: Object-Oriented Programming and Classes (individual)

In this assignment, you should write your code in a **readable** way, and **modularise** chunks of code that would otherwise be repeated.

Your function and class definitions should have appropriate **docstrings**.

You are expected to follow PEP8 conventions as far as possible. Run the code cell below to load a Jupyter extension that will inform you of any non-compliant code.

```python
# This code cell loads a PEP8 linter.
# Linting is the process of flagging programming errors,
# bugs, stylistic errors, and other code problems
# pycodestyle is a linter that highlights any syntax that
# is not PEP8-compliant.
# You only need to load it once in each notebook.

%load_ext pycodestyle_magic
%pycodestyle_on
```

## Part 1: Implement a `Temperature` class

Temperature conversion is such a pain. The same temperature can be represented in units of °C, F, or K! It would be really handy if we could bundle them up into a `Temperature` class that will make it easy for us to think about temperature.

### Fomulas

**Celsius (C) to Kelvin (K)**  
$K = C + 273.15$

**Celsius (C) to Fahrenheit (F)**  
$F = (C\times\frac{9}{5})+32$

### Task 1: define a `Temperature` class

1. Define an `__init__()` method that takes in a temperature (in °C), and stores three **attributes**:
   - `celsius`, which returns the temperature in degrees Celsius
   - `fahrenheit`, which returns the temperature in Fahrenheit
   - `kelvin`, which returns the temperature in Kelvin


2. The method should raise a `TypeError` if the input is not `int` or `float` type.


3. The method should raise a `ValueError` if the temperature is invalid (e.g. below absolute zero).
   - C >= –273.15 °C


4. Define a `__repr__()` method that returns the expression used to initialise the object
   - e.g. Temperature(0)



<!-- grade_id: cell-261b0ff2b9c665c0, points: 0, grade: false, solution: true -->
```python
# Define a Temperature class


class Temperature:
    '''
    A class that stores a temperature value and facilitates conversion
    between units.
    '''
    def __init__(self, temp_c):
        ### Type your code below
        
```

<!-- grade_id: cell-4e8031a164b15436, points: 2, grade: true, solution: false -->
```python
# AUTOGRADING: Formula correctness check
test_temp = Temperature(0)
assert test_temp.celsius == 0, \
    ('Wrong value for celsius attribute, '
     f'got {test_temp.celsius} instead of 0.')
assert test_temp.fahrenheit == 32, \
    ('Wrong value for fahrenheit attribute, '
     f'got {test_temp.fahrenheit} instead of 32.')
assert test_temp.kelvin == 273.15, \
    ('Wrong value for kelvin attribute, '
     f'got {test_temp.kelvin} instead of 273.15.')

test_temp = Temperature(100)
assert test_temp.celsius == 100, \
    (f'Wrong value for celsius attribute, '
     f'got {test_temp.celsius} instead of 100.')
assert test_temp.fahrenheit == 212, \
    (f'Wrong value for fahrenheit attribute, '
     f'got {test_temp.fahrenheit} instead of 212')
assert test_temp.kelvin == 373.15, \
    (f'Wrong value for kelvin attribute, '
     f'got {test_temp.kelvin} instead of 373.15.')
```

<!-- grade_id: cell-3e66e69d63d925af, points: 1, grade: true, solution: false -->
```python
# AUTOGRADING: Type check
import unittest


class TypeCheck(unittest.TestCase):
    def testTypeError(self):
        self.assertRaises(TypeError, Temperature, '38°C')
        self.assertRaises(TypeError, Temperature, [38])


# unittest.main looks at sys.argv and first parameter is what
# started IPython or Jupyter
# exit=False prevents unittest.main from shutting down the kernel process
result = unittest.main(argv=['ignored'], exit=False)
```

<!-- grade_id: cell-94d7520a92dac8f4, points: 1, grade: true, solution: false -->
```python
# AUTOGRADING: Value check
import unittest


class TypeCheck(unittest.TestCase):
    def testTypeError(self):
        self.assertRaises(ValueError, Temperature, -273.16)
        self.assertRaises(ValueError, Temperature, -274)


result = unittest.main(argv=['ignored'], exit=False)
```

### Task 2a: Define a `set()` method to change the value of temperature

Now let’s have a way to change the initial value of temperature; let’s add a `set()` method to `Temperature`.

1. Define a `set()` method that takes in an argument `value` (`float` or `int`) representing the temperature in units of °C.
2. `set()` should update the value of temperature in Fahrenheit and Kelvin as well.

### Task 2b: Update the `__init__()` method

Now it looks like `__init__()` and `set()` perform a very similar task. Instead of repeating code, let’s update `__init__()` to call `set()` instead.

3. Update `__init__()` to call `set()` instead.

<!-- grade_id: cell-02abfd3e847b0f1e, points: 3, grade: true, solution: true -->
```python
# Define a set() method
# Update the __init__() method


class Temperature:
    '''
    A class that stores a temperature value and facilitates conversion
    between units.
    '''
    def __init__(self, temp_c):
        ### Type your code below
        
```

We won't always get our input in units of degree Celsius; sometimes we may get input in Fahrenheit or Kelvin. We need a way to allow `__init__()` to accept those units as well.

### Task 3: update `set()` to accept other units

1. Update the `__init__()` method to take in two arguments:
   - a temperature value
   - a keyword argument, `units=` defining the units of the temperature value
       - 'c' for degrees Celsius (default)
       - 'k' for Kelvin
       - 'f' for Fahrenheit
       
   - The method should still meet the other requirements of **Task 1**.
   - You should also update `__init__()` and `__repr__()` to follow the new syntax of `set()`.

<!-- grade_id: cell-411419efe239e9d4, points: 0, grade: false, solution: true -->
```python
# Update the set() method


class Temperature:
    '''
    A class that stores a temperature value and facilitates conversion
    between units.
    '''
    def __init__(self, temp, **kwargs):
        # Paste your code from earlier
        self.set(temp, units=kwargs.get('units', 'c').lower())

    def __repr__(self):
        # Paste your code from earlier
        pass

    def set(self, temp, **kwargs):
        '''
        Updates the temperature value, depending on the units given.

        Usage:
        temp.set(temp, units='c')
            sets the temperature to be temp in degrees Celsius

        temp.set(temp, units='f')
            sets the temperature to be temp in degrees Fahrenheit

        temp.set(temp, units='k')
            sets the temperature to be temp in Kelvin
        '''
        # Validation code is already written for you here
        units = kwargs.get('units', 'c').lower()
        if units is None:
            raise KeyError('No units specified '
                           '(must be \'c\', \'f\', or \'k\')')
        elif type(units) != str:
            raise TypeError(f'Invalid input type: got {type(units)}, '
                            'str expected')
        elif units not in 'cfk':
            raise KeyError('Invalid units specified '
                           '(must be \'c\', \'f\', or \'k\')')

        if type(temp) not in (int, float):
            raise TypeError('Invalid input type, '
                            'must be int or float.')
        ### Type your code here
        # No return statement needed
        
```

<!-- grade_id: cell-874e8166cd7d6bcf, points: 3, grade: true, solution: false -->
```python
# AUTOGRADING: Value check

test_temp = Temperature(0, units='c')
assert test_temp.celsius == 0, \
    ('Wrong value for celsius attribute, '
     f'got {test_temp.celsius} instead of 0.')
assert test_temp.fahrenheit == 32, \
    ('Wrong value for fahrenheit attribute, '
     f'got {test_temp.fahrenheit} instead of 32.')
assert test_temp.kelvin == 273.15, \
    ('Wrong value for kelvin attribute, '
     f'got {test_temp.kelvin} instead of 273.15.')

test_temp = Temperature(0, units='f')
assert abs(test_temp.celsius - (-17.8)) < 0.05, \
    ('Wrong value for celsius attribute, '
     f'got {test_temp.celsius} instead of -32.')
assert test_temp.fahrenheit == 0, \
    ('Wrong value for fahrenheit attribute, '
     f'got {test_temp.fahrenheit} instead of 0.')
assert abs(test_temp.kelvin - 255.4) < 0.05, \
    ('Wrong value for kelvin attribute, '
     f'got {test_temp.kelvin} instead of 255.4.')

test_temp = Temperature(0, units='k')
assert test_temp.celsius == -273.15, \
    ('Wrong value for celsius attribute, '
     f'got {test_temp.celsius} instead of -273.15.')
assert abs(test_temp.fahrenheit - (-459.7)) < 0.05, \
    ('Wrong value for fahrenheit attribute, '
     f'got {test_temp.fahrenheit} which is too far from -459.7.')
assert test_temp.kelvin == 0, \
    ('Wrong value for kelvin attribute, '
     f'got {test_temp.kelvin} instead of 0.')
```

# Feedback and suggestions

Any feedback or suggestions for this assignment?

<!-- grade_id: cell-63f8f1cb7ff65bb1, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE