## 1-minute introduction to Jupyter ##

A Jupyter notebook consists of cells. Each cell contains either text or code.

A text cell will not have any text to the left of the cell. A code cell has `In [ ]:` to the left of the cell.

If the cell contains code, you can edit it. Press <kbd>Enter</kbd> to edit the selected cell. While editing the code, press <kbd>Enter</kbd> to create a new line, or <kbd>Shift</kbd>+<kbd>Enter</kbd> to run the code. If you are not editing the code, select a cell and press <kbd>Ctrl</kbd>+<kbd>Enter</kbd> to run the code.

---

# Lesson 10a: Advanced functions

In lesson 5, we learnt one way of writing modules of code that can be reused: by defining a function. We learnt that a function takes in **parameters**, and **returns** an output:

    def isfloat(string):
        if ('.' in string) and (string.count('.') == 1):
            return True
        else:
            return False

In the above code snippet, the defined function `isfloat()` takes in a parameter<sup>[1]</sup>, `string`, and returns an output value (either `True` or `False`).

    >>> to_hms(7199,'list')
    [1,59,59]
    >>> to_hms(7199,'string')
    '1 hour, 59 minutes, 59 seconds'

In the above code snippet, the called function `to_hms()` takes in two arguments<sup>[1]</sup>, `7199` and `'list'` and returns a list. But if it takes in arguments `7199` and `'string'`, it returns a string.

[1]: **See:** [What's the difference between a parameter and an argument?](#What's-the-difference-between-a-parameter-and-an-argument?)

These arguments are known as **positional arguments**, because their position matters. I cannot call the function as `to_hms('list',7199)`; the function will assign the arguments `7199` and `'list'` to the wrong parameters in the definition!

How do we write even more complex functions?

## The weirdness of Python's `print()` function

You might have noticed by now that Python's `print()` function is really strange ... and really powerful.

`print()` needs a `str` input to print to the console; any non`str` types must first be casted to `str`. So it can't do this:

    >>> print('1 + 2 = '+(1+2))
    TypeError: must be str, not int
    
But it can do this:

    >>> print(f'1 + 2 = {1 + 2}')
    1 + 2 = 3
    >>> print('1 + 2 =', 1 + 2)
    1 + 2 = 3
    
We know that `print(f'1 + 2 = {1 + 2}')` is making use of an f-string which will auto-cast non-`str` types to `str`.

But what is going on with `print('1 + 2 =', 1 + 2)`? `print()` can take 1 argument, but it can also take 2 arguments?

In fact, this works too:

    >>> print('1 + 2', '=', 1 + 2)
    1 + 2 = 3
    
And so does this:

    >>> print('1', '+', '2', '=', 1 + 2)
    1 + 2 = 3
    
How is a function able to accept 1 argument, 2 arguments, 3 arguments, and even 5 arguments?

## Using multiple arguments (args) in a function

Let's extend our `isfloat()` function to accept multiple arguments and return a value (`True` or `False`) for each argument. We want it to be able to do this:

    >>> a, b, c = isfloat('1.0', '200', '.3')
    >>> a
    True
    >>> b
    False
    >>> c
    True
    
And we want this function to be able to handle any number of arguments (more than 0).

Let's start with a simple function definition:

```python
def isfloat(*args):
    '''Write the rest of the function here.'''
    # This function does nothing
    pass
```

### The `pass` keyword: Python's way of doing nothing

`pass` is a special Python keyword that does ... nothing. With it, we can declare the functions we need first, use `pass` as a placeholder for code we will write in future, and plan out the rest of the program while writing minimal code.

We need to `pass` keyword because Python won't let us write any statements ending in a colon(`:`) without any code after it. Likewise, we cannot write `if .. else` statements that have no code after the statement.

This means you can't use `pass` as a variable name ... important to know when using variable names to store user passwords.

Notice that in this function declaration, instead of using a positional argument `string` like we did at the start, we use another argument, called `*args`. What kind of variable is that?

Let's write some code to investigate what this `*args` is. We'll simply make the function return `args` so we can inspect it:

```python
def return_args(*args):
    '''Let's investigate what args is.'''
    # Let's return args as the function output so we can inspect it
    # after calling the function
    return args

args = return_args('1.0', '200', '.3')
args
```

Interesting ... Python bundles up all the arguments we fed to `isfloat()` and assigns it to the `args` parameter as a tuple!

Actually, we don't need to name it `args`, any other name will do as well. `args` is a common name used by convention so that other Python programmers an identify it easily:

```python
def isfloat(*anyvariablenamewilldo):
    '''
    To show that any variable name will do for the parameter.
    The important thing is to make sure there is only one asterisk
    in front of it.
    '''
    return anyvariablenamewilldo

args = isfloat('1.0', '200', '.3')
args
```

So now we can iterate over this tuple and process it:

```python
def isfloat(*args):
    '''
    Takes in multiple arguments and returns a boolean value for each.
    True if the string is a float, False if it is not.
    '''
    result = [] # Initialise a list instead of tuple to store the results
                # since tuples are immutable
    for string in args: # Iterate over args tuple
        if ('.' in string) and (string.count('.') == 1):
            result.append(True)
        else:
            result.append(False)
    return tuple(result) # Convert result to tuple before returning

isfloat('1.0','200','.3')
```

### Returning immutable results as a tuple

You saw earlier that Python bundles up the arguments into `args` and returns it as a tuple, not as a list; we try to be consistent with this practice in our code.

In general, we try to return collections of immutable results as tuples and not lists, unless the output is intended to be a list. This prevents any accidental modification of the result.

## Using positional arguments with multiple arguments

Let's try to extend multiple argument support to other functions we have written, such as `round_dp(n, num_str)` which rounds a given numeral string `num_str` to the requested number of decimal places `dp`:

```python
def round_dp(n, num_str):
    '''
    Usage:
    round_dp(num_str,n)
    
    Round an input numerical string num_str to n number of decimal places.
    '''
    # We are simplifying this function by skipping the usual validation checks
    # We assume there is always a decimal place,
    # the number of dp is >= 0,
    # and the requested number of dp will always be smaller than dp of number
    decimal_position = num_str.find('.')
    stripped_num_str = num-str.replace('.', '')
    check_digit_position = decimal_position + n
    check_digit = int(num_str[check_digit_position])
    num_to_round = int(num_str[:check_digit_position])
    if check_digit >= 5:
        num_to_round += 1
    num_to_round = str(num_to_round)

    return f'{num_to_round[:decimal_position]}.{num_to_round[decimal_position:]}'
```

How do we get this function to support multiple arguments? We will always need  the positional argument `n`, and at least one input argument to `args`.

One way is to simply extract `n` from the `args` tuple:

```python
def round_dp(*args):
    n = args[0] # assign first argument to n
    del args[0] # remove first argument from args
    # args should now contain the remaining arguments
    [...]
```

But this method is prone to failure if the user does not use the function as intended. It is better for Python to be aware that the first argument is positional, while the remaining arguments can be grouped into a (tuple) collection.

We do that by declaring positional arguments first, and multiple arguments later:

```python
def round_dp(n, *args):
    # Skipping docstring and other comments for brevity
    result = [] # added line
    for num_str in args:  # added line
        decimal_position = num_str.find('.')
        stripped_num_str = num_str.replace('.', '')
        check_digit_position = decimal_position + n
        check_digit = int(num_str[check_digit_position])
        num_to_round = int(num_str[:check_digit_position])
        if check_digit >= 5:
            num_to_round += 1
        num_to_round = str(num_to_round)
        result.append(f'{num_to_round[:decimal_position]}.{num_to_round[decimal_position:]}')  # added line
    return result  # added line
```

We can use as many positional arguments as we like, provided they are declared **before \*args**.

Q1: Can positional arguments with default values be mixed with multiple arguments? If yes, how?

A1: Yes. Parameters with default values must be defined after parameters without default values.

## Multipurpose functions

You have just seen how to make a function take in an arbitrary number of arguments. And you already know how to make a function use default values for unspecified arguments.

But what if we have to write functions that are really advanced and have all kinds of parameters?

For instance, the `print()` function lets you set the separator (default: `' '`):

    >>> print('1', '+', '1', '=', '2', sep='^')
    1^+^1^=^2
    
It also lets you decide what character to print at the end of each line (default: `\n`)

    >>> print('1', '+', '1', '=', '2', end=';')
    1 + 1 = 2;
    
And we can also combine both options:

    >>> print('1', '+', '1', '=', '2', sep='^', end=';')
    1^+^1^=^2;
    
Notice that these argument require a different specification: we have to call them by name. These arguments are known as **keyword arguments**.

## Using keyword arguments in a function

Let’s write a test function to help us examine how keyword arguments are used.

We declare the function this way:

```python
def return_kwargs(**kwargs):
    '''Let's investigate kwargs.'''
    # Let's just return kwargs and see what we get.
    return kwargs

kwargs = return_kwargs('1.234', sf=2) # '1.234' will be stored in args
kwargs
# Try calling return_kwargs with dp=2 and observe the result
# Try calling return_kwargs with sf=2, dp=2 and observe the result
```

The keyword arguments are stored in `kwargs` and returned as a dictionary. So simple!

Again, naming it `kwargs` is just a convention; you can name the parameter anything you like, so long as there are two asterisks in front of it.

Access the arguments through the `kwarg` dictionary like any ordinary dictionary:

```python
# Call return_kwargs above with an `sf=` keyword argument first.
# The run this cell to observe the return value of kwargs['sf']
kwargs['sf']
```

If no `sf=` keyword argument was given, that would throw a `KeyError` since `kwargs` would not have a `'sf'` key. To get around this, we can do a membership check with the `in` keyword:

    def return_kwargs(**kwargs):
        if 'sf' in kwargs.keys():
            sf = kwargs['sf']
        
or we can use the `dict.get()` method to [assign a fallback value](lesson_07a.ipynb) if the key is not present. This chunk of code will assign the value `None` to the `sf` parameter if `return_kwargs()` was called without a `sf=` keyword argument:

    def return_kwargs(**kwargs):
        sf = kwargs.get('sf',None)
        if sf is not None:
            [... the rest of your code here ...]

We will use these examples shortly. But first, let’s learn more advanced ways of handling errors: by raising the right errors where relevant.

### Raising an error with the `raise` statement

To raise an error, use the `raise` keyword:

```python
raise SyntaxError
```

To provide more information, you can also provide a message as an argument to the error:

```python
raise SyntaxError('Playing with the raise keyword')
```

Great. Let’s get back to using keyword arguments and raising errors.

## Combining `round_dp()` and `round_sf()` functions

We wrote two functions in assignment 5: `round_dp()` and `round_sf()`. It would be better if we had one function, `round_to()` that could handle both dp and sf arguments, and round accordingly.

How should we define this function?

    def round_to(num_str, sf_or_dp):
    
Will this work? We wouldn't know if the second argument was meant to be `dp` or `sf`, so this would not work. How about this:

    def round_to(num_str, sf=None, dp=None):
    
Ehhh … it's possible, but if you want to give a dp argument you would have to do something like `round_to('123.456', None, 2)` to round `'123.456'` to 2 dp. That’s not nice. It’s not _elegant_.

It would be nice if we could call `round_to('123.456', dp=2)` to round to 2 dp, and `round_to('123.456', sf=2)` to round to 2 sf. It’s easy to read, with no ambiguity. This means we should use keyword arguments to implement this function.

But first, we have to do some validation. The user must give _either_ the `sf=` or the `dp=` argument, but not both together. Let’s focus on the validation first without worrying about the rest of the function:

### Task 1: Prevent two keyword arguments being used at the same time

Complete the function definition below to validate the keyword arguments and:

1. raise a `TypeError` if either the `sf=` or `dp=` keyword arguments are non-integers, and provide a helpful error message;
2. raise a `ValueError` if the `sf=` and `dp=` keyword arguments are used together
3. raise a `TypeError` if the `sf=` and `dp=` keyword arguments are not provided at all

<!-- grade_id: cell-6288888c35f967c4, points: 0, grade: false, solution: true -->
```python
# Step 1 is done for you below.
# Complete the code below to complete validation steps 2 and 3
def round_to(num_str,**kwargs):
    '''Validate the input to keyword arguments.'''
    if 'sf' in kwargs.keys() and 'dp' in kwargs.keys():
        raise ValueError('Cannot use dp= and sf= keyword arguments at the same time.')
    # Continue your code here
    
```

<!-- grade_id: cell-ea2519e82a89b476, points: 3, grade: true, solution: false -->
```python
# Test cell to check your code above
# We can't use assert statements to check if errors are raised, since those errors will halt the program.
# We will need to use other Python libraries, such as unittest.
# This will not be covered in the syllabus, but you can read up on it
# at https://docs.python.org/3.6/library/unittest.html
# You will see 'OK' at the end of the output if all tests passed
import unittest

class ErrorCheck(unittest.TestCase):
    def testKwargTypeError(self):
        self.assertRaises(TypeError,round_to, '123.456', sf='2')
        self.assertRaises(TypeError,round_to, '123.456', dp='2')

    def testKwargNumError(self):
        self.assertRaises(ValueError,round_to, '123.456', dp=2, sf=2)

    def testKwargNoneError(self):
        self.assertRaises(KeyError,round_to, '123.456')

result = unittest.main(argv=['ignored'], exit=False)
```

Okay, you’ve done some basic validation for the keyword arguments. Let’s write the rest of the function, assuming that the input is valid. We will focus on the core features of the rounding, without worrying about negative numbers and other edge cases (e.g. 99.999).

Let's look at `round_dp()` and `round_sf()` again. Here, they have been rewritten and simplified to highlight their similarities:

```python
round_dp(n,num_str):
    if '.' in num_str:
        decimal = num_str.find('.') # decimal position
        num_str = num_str.replace('.', '') # remove decimal
    else:
        decimal = None
    check_pos = decimal + n # position of check digit
    check_digit = int(num_str[check_pos])
    num_to_round = int(num_str[:check_pos])
    if check_digit >= 5:
        num_to_round += 1
    rounded_num = str(num_to_round)
    if decimal is not None and decimal < len(rounded_num):
        return f'{rounded_num[:decimal]}.{rounded_num[decimal:]}' # added line
    else:
        return rounded_num
```

```python
round_sf(n,num_str):
    if '.' in num_str:
        decimal = num_str.find('.') # decimal position
        num_str = num_str.replace('.', '') # remove decimal
    else:
        decimal = None
    check_pos = n # position of check digit
    check_digit = int(num_str[check_pos])
    num_to_round = int(num_str[:check_pos])
    if check_digit >= 5:
        num_to_round += 1
    rounded_num = str(num_to_round)
    if decimal is not None and decimal < len(rounded_num):
        return f'{rounded_num[:decimal]}.{rounded_num[decimal:]}' # added line
    else:
        return rounded_num
```

## Exercise 1: Refactor `round_dp()` and `round_sf()`into `round_to()`

### Task 2

Combine `round_dp()` and `round_sf()` as defined above into one function, `round_to()`, that can take either a `dp=` or a `sf=` declaration.

This act of rewriting code to support more advanced features, or just to simplify the code, without changing its functionality, is known as [**refactoring**](https://refactoring.com/). You will see this term _a lot_ in the software industry.

**Example output**

    >>> round_to('1.234', dp=2)
    '1.23'
    >>> round_to('1.234', sf=2)
    '1.2'
    
Let’s ignore the validation for now, so you can focus on refactoring `round_dp()` and `round_sf()`.

<!-- grade_id: cell-cc28d8a35209cf7e, points: 0, grade: false, solution: true -->
```python
def round_to(num_str, **kwargs):
    '''
    Usage:
    round_to(num_str, [sf=int][, dp=int])
    
    Rounds a given numerical string to the desired number of dp or sf.
    Either dp or sf may be specified, but not both.
    '''
    # Refactor round_dp() and round_sf().
    # Ignore validation steps for now.
    # Write your code below.
    
```

<!-- grade_id: cell-0c41a018dfae6c71, points: 2, grade: true, solution: false -->
```python
# Test cell to check your code above
# We can't use assert statements to check if errors are raised, since those errors will halt the program.
# We will need to use other Python libraries, such as unittest.
# This will not be covered in the syllabus, but you can read up on it
# at https://docs.python.org/3.6/library/unittest.html
# You will see 'OK' at the end of the output if all tests passed
import unittest

class ErrorCheck(unittest.TestCase):
    def testRoundToDp(self):
        testdata = [('1.23456', 0, '1'),
                    ('1.23456', 1, '1.2'),
                    ('1.23456', 2, '1.23'),
                    ('1.23456', 3, '1.235'),
                    ('1.23456', 4, '1.2346'),
                    ('1.23456', 5, '1.23456'),
                    ('-1.2345', 3, '-1.234'),
                    ]
        for num, dp, ans in testdata:
            result = round_to(num, dp=dp)
            self.assertEqual(result, ans)

    def testRoundToSf(self):
        testdata = [('1.23456', 1, '1'),
                    ('1.23456', 2, '1.2'),
                    ('1.23456', 3, '1.23'),
                    ('1.23456', 4, '1.235'),
                    ('1.23456', 5, '1.2346'),
                    ('1.23456', 6, '1.23456'),
                    ('-1.23456', 4, '-1.234'),
                    ('123456', 4, '123500'),
                    ]
        for num, sf, ans in testdata:
            result = round_to(num, sf=sf)
            self.assertEqual(result, ans)

result = unittest.main(argv=['ignored'], exit=False)
```

Okay, now you can add in validation.

### Task 3: Validate the input for `round_to()`

By combining your code from Task 1 with the function above, raise an appropriate error when invalid keyword arguments are provided.

Also write an appropriate docstring for the function.

You are not expected to get `round_to()` working perfectly for this lesson. If you can implement keyword arguments and raise errors appropriately, you have achieved the lesson objective.



**Example output**

    >>> round_to('1.234', dp=2, sf=2)
    ValueError
    >>> round_to('1.234')
    ValueError
    >>> round_to('1.234', sf='0')
    TypeError

<!-- grade_id: cell-391a1dfb7eee168d, points: 0, grade: false, solution: true -->
```python
def round_to(num_str, **kwargs):
    # (1) Write a docstring
    # (2) Validate the keyword arguments.
    # (3) Round the function to the requested number of dp/sf.
    
```

<!-- grade_id: cell-021b53d1d1ebd631, points: 3, grade: true, solution: false -->
```python
# Test cell to check your code above
# We can't use assert statements to check if errors are raised, since those errors will halt the program.
# We will need to use other Python libraries, such as unittest.
# This will not be covered in the syllabus, but you can read up on it
# at https://docs.python.org/3.6/library/unittest.html
# You will see 'OK' at the end of the output if all tests passed
import unittest

class ErrorCheck(unittest.TestCase):
    def testKwargTypeError(self):
        self.assertRaises(TypeError, round_to, '123.456', sf='2')
        self.assertRaises(TypeError, round_to, '123.456', dp='2')

    def testKwargNumError(self):
        self.assertRaises(ValueError, round_to, '123.456', dp=2, sf=2)

    def testKwargNoneError(self):
        self.assertRaises(TypeError, round_to, '123.456')
        
    def testRoundToDp(self):
        testdata = [
            ('1.23456', 0, '1'),
            ('1.23456', 1, '1.2'),
            ('1.23456', 2, '1.23'),
            ('1.23456', 3, '1.235'),
            ('1.23456', 4, '1.2346'),
            ('1.23456', 5, '1.23456'),
            ('-1.2345', 3, '-1.234'),
        ]
        for num, dp, ans in testdata:
            result = round_to(num, dp=dp)
            self.assertEqual(result, ans)

    def testRoundToSf(self):
        testdata = [
            ('1.23456', 1, '1'),
            ('1.23456', 2, '1.2'),
            ('1.23456', 3, '1.23'),
            ('1.23456', 4, '1.235'),
            ('1.23456', 5, '1.2346'),
            ('1.23456', 6, '1.23456'),
            ('-1.23456',4, '-1.234'),
            ('123456', 4, '123500'),
        ]
        for num, sf, ans in testdata:
            result = round_to(num, sf=sf)
            self.assertEqual(result, ans)

result = unittest.main(argv=['ignored'], exit=False)
```

Was it tedious testing your `round_to()` function? I bet! We will learn more efficient ways of testing your code in the next lesson. This code testing is _much easier_ if you learn to write your code as functions instead of procedures.