## 1-minute introduction to Jupyter ##

A Jupyter notebook consists of cells. Each cell contains either text or code.

A text cell will not have any text to the left of the cell. A code cell has `In [ ]:` to the left of the cell.

If the cell contains code, you can edit it. Press <kbd>Enter</kbd> to edit the selected cell. While editing the code, press <kbd>Enter</kbd> to create a new line, or <kbd>Shift</kbd>+<kbd>Enter</kbd> to run the code. If you are not editing the code, select a cell and press <kbd>Ctrl</kbd>+<kbd>Enter</kbd> to run the code.

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Assignment 10: Advanced modular programming, testing

In this assignment, you should write your code in a **readable** way, and **modularise** chunks of code that would otherwise be repeated.

Modularising a section of program code means to reorganise it in a way that allows it to be reused in another part of the code easily, usually in the form of a function.

Your function definitions should have appropriate **docstrings**.

## Part 1: Testing a Sudoku solver

A 3×3 sudoku solver receives the following generated input:

```python
puzzle = [[None, 9, None],
          [3, None, 7],
          [None, 1, None],
          ]
```

It is called to solve the puzzle, and it returns the following output:

```python
solution = [[4, 9, 2],
            [3, 5, 7],
            [8, 1, 6],
            ]
```

**Task:** Write `assert` statements to _verify_ that the solution is correct.

A correct solution should have all its rows, columns, and diagonals sum to 15.

<!-- grade_id: cell-e03c120171de668c, points: 3, grade: true, solution: true -->
```python

```

## Part 2: Setting a `debug` flag for `print`ing

Now that you have learned how to use keyword arguments, let’s try using them for something potentially useful: `print()` statements for debugging.

When you don’t need them, you are afraid to remove them, because re-inserting them again is so annoying. Is there a better alternative to commenting out multiple `print()` statements, then uncommenting them again?

Yes! Let’s use keyword arguments to determine whether the function should print debug statements or not.

A sample `text_numeral()` function is provided in the code cell below.

### Task: Add keyword arguments

Edit the code so that the `print()` statements are executed only when the following function call is made:

    >>> text_numeral(367, debug=True)
    prev_key: 0
    this_key: 1
    prev_key: 1
    [...] <-- represents truncated output
    selected key: 7
    output: three hundred and sixty-seven
    remainder: 0
    'three hundred and sixty-seven' <-- this is the output value (str type)

but not when it is called without the `debug` keyword argument:

    >>> text_numeral(367)
    'three hundred and sixty-seven' <-- this is the output value (str type)

**Hint:** The keyword arguments are passed to the function as a dictionary. you can use the dict’s `.get()` method to retrieve the keyword argument, and set it to a fallback value if it is not provided.

<!-- grade_id: cell-06d02db7c32d0c73, points: 0, grade: false, solution: true -->
```python
# Edit this code to print() only when debug=True is passed to the function as a keyword argument

def text_numeral(num:int, **kwargs):
    '''
    Takes in a number below 1000 and converts it to text form.
    
    Usage:
    text_numeral(num)
    '''
    
    
    if num >= 1000:
        raise ValueError('input must be less than 1000.')

    numeral = {1: 'one', 2: 'two', 3: 'three', 4: 'four', 5: 'five',
               6: 'six', 7: 'seven', 8: 'eight', 9: 'nine', 10: 'ten',
               11: 'eleven', 12: 'twelve', 13: 'thirteen', 14: 'fourteen', 15: 'fifteen',
               16: 'sixteen', 17: 'seventeen', 18: 'eighteen', 19: 'nineteen', 20: 'twenty',
               30: 'thirty',
               40: 'forty',
               50: 'fifty',
               60: 'sixty',
               70: 'seventy',
               80: 'eighty',
               90: 'ninety',
               }

    output = ''
    while num > 0:
        # Select largest key that is still smaller than num
        # Since we need to check that this_key is greater than num,
        # we need to store prev_key and prev_word since they are the
        # actual key and value we want.
        key, word = 100, 'hundred'
        prev_key, prev_word = 0, None
        for this_key, this_word in numeral.items():
            print(f'prev_key: {prev_key}')
            print(f'this_key: {this_key}')
            if prev_key <= num < this_key:
                print(f'match: {prev_key} <= {num} < {this_key}')
                key, word = prev_key, prev_word
            prev_key, prev_word = this_key, this_word
        print(f'selected key: {key}')

        # Build up the output string
        if key == 100:
            hundred_count, num = divmod(num, key)
            output += f'{numeral[hundred_count]} {word}'
            if num > 0:
                output += ' and '
        else:
            output += word
            num -= key
            if num > 0:
                output += '-'
        print(f'output: {output}')
        print(f'remainder: {num}')
        if num == 0:
            return output
```

<!-- grade_id: cell-0d323b2b11660e53, points: 2, grade: true, solution: false -->
```python
# Autograding test for no debug kwarg
# If your code works, this cell should produce no output when run


import sys
from io import StringIO


capture = StringIO()
temp = sys.stdout
sys.stdout = capture
result = text_numeral(367)
sys.stdout = temp
assert capture.getvalue().strip() == '', 'Function printed output when it should be silent'
assert result == 'three hundred and sixty-seven', f'output from function \'{result}\' is incorrect'
```

<!-- grade_id: cell-1c657f7fe577e4be, points: 2, grade: true, solution: false -->
```python
# Autograding test for debug=True kwarg
# If your code works, this cell should produce no output when run


import sys
from io import StringIO


capture = StringIO()
temp = sys.stdout
sys.stdout = capture
result = text_numeral(367, debug=True)
sys.stdout = temp
for each in ('selected key: 7',
             'output: three hundred and sixty-seven',
             'remainder: 0',
             ):
    assert each in capture.getvalue().strip(), f'expected print statement not found: {each}'
assert result == 'three hundred and sixty-seven', f'output from function \'{result}\' is incorrect'
```

## Part 2: Write a `try-except` structure

The Open Library API (application programming interface) offers programmers a way to search for books and retrieve book information in their program. The following functions are provided below:

1. `search()` takes in a `searchterm` (`str`) argument and returns search results from the Open Library API
2. `save_to_csv()` takes in `list_of_dict` (`list`) and `filepath` (`str`) positional arguments, and writes the title, author name, and first year of publication of the results to a CSV file located in `filepath`.

```python
# Run this cell to use the functions below
# Do not modify the functions

from openlibrary import search, save_to_csv
```

```python
help(search)
```

```python
help(save_to_csv)
```

A developer is trying to use these two functions to save search results to a CSV file in _his_ function `search_and_save()` (in the code cell below). He wants the search results to be saved into a folder representing the current date, to keep the files organised. But he keeps encountering a `FileNotFoundError`.

This `FileNotFoundError` happens because the directory `'20200420'` does not exist yet. He needs to call `os.mkdir(folder)` to create the folder. But `os.mkdir()` will raise a `FileExistsError` if the directory already exists.

**Task:** Modify the code in `search_and_save()` below by adding `try-except` statements so that these errors are handled successfully.

**Hint:** You may have to write nested `try-except` statements to handle further Exceptions that occur in the `except`, `else`, or `finally` statements

<!-- grade_id: cell-89d8ab25ade8cd31, points: 0, grade: false, solution: true -->
```python
# Modify the code in this cell with try-except-else statements
# To take appropriate actions when exceptions occur
# And write the CSV file successfully

import os
from datetime import date


def search_and_save(searchterm):
    '''
    Search for results from Open Library API for searchterm,
    and print the results to a file in a folder.
    
    Folder is named by date, YYYYMMDD format.
    
    File path: ./<YYYYMMDD>/<searchterm>.csv
    '''
    result = search(searchterm)
    folder = date.today().strftime('%Y%m%d')
    filepath = f'./{folder}/{searchterm}.csv'
    save_to_csv(result['docs'], filepath)
    
```

<!-- grade_id: cell-77b553ad9dd7a273, points: 4, grade: true, solution: false -->
```python
# AUTOGRADING: Test for file creation and correct data

import os, shutil
from datetime import date
import unittest
from pathlib import Path


def readfile(filepath):
    if os.path.exists(filepath):
        with open(filepath, 'r') as f:
            data = [row.strip().split(',') for row in f.readlines()]
    return data.pop(0), data

def remove_notexist_raise(filepath):
    try:
        assert os.path.exists(filepath), f'File not found: {filepath}'
    except Exception:
        raise
    else:
        shutil.rmtree(Path(filepath).parent)


class testCheck(unittest.TestCase):
    def testNotRaiseError(self):
        searchterm = '50 Shades of Grey'
        try:
            search_and_save(searchterm)
        except Exception as e:
            print(f'{e.__class__} Test failed, exception raised: {e}')
        finally:
            folder = date.today().strftime('%Y%m%d')
            filepath_to_check = f'./{folder}/{searchterm}.csv'
            if os.path.exists(filepath_to_check):
                shutil.rmtree(f'./{folder}')

    def testFileCreated(self):
        searchterm = 'Harry Potter'
        folder = date.today().strftime('%Y%m%d')
        filepath = f'./{folder}/{searchterm}.csv'
        try:
            search_and_save(searchterm)
        except Exception as e:
            print(f'{e.__class__} Test failed, exception raised: {e}')
        else:
            assert os.path.isfile(filepath), f'File not found: {filepath}'
        finally:
            remove_notexist_raise(filepath)
        
    def testDataHeader(self):
        searchterm = 'Harry Potter'
        folder = date.today().strftime('%Y%m%d')
        filepath = f'./{folder}/{searchterm}.csv'
        try:
            search_and_save(searchterm)
        except Exception as e:
            print(f'{e.__class__} Test failed, exception raised: {e}')
        else:
            header, data = readfile(filepath)
            assert header == ['title', 'author_name', 'first_publish_year'], \
                f'Wrong header ({header} != [\'title\', \'author_name\', \'first_publish_year\'])'
        finally:
            remove_notexist_raise(filepath)

    def testDataFormat(self):
        searchterm = 'Harry Potter'
        folder = date.today().strftime('%Y%m%d')
        filepath = f'./{folder}/{searchterm}.csv'
        try:
            search_and_save(searchterm)
        except Exception as e:
            print(f'{e.__class__} Test failed, exception raised: {e}')
        else:
            header, data = readfile(filepath)
            for row in data:
                assert len(row) == 3, f'Wrong number of elements in row ({len(row)}), expected 3\n' + \
                                      f'Row data: {row}'
        finally:
            remove_notexist_raise(filepath)

# unittest.main looks at sys.argv and first parameter is what started IPython or Jupyter
# exit=False prevents unittest.main from shutting down the kernel process
result = unittest.main(argv=['ignored'], exit=False)
```

# Feedback and suggestions

Any feedback or suggestions for this assignment?

<!-- grade_id: cell-63f8f1cb7ff65bb1, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE