## 1-minute introduction to Jupyter ##

A Jupyter notebook consists of cells. Each cell contains either text or code.

A text cell will not have any text to the left of the cell. A code cell has `In [ ]:` to the left of the cell.

If the cell contains code, you can edit it. Press <kbd>Enter</kbd> to edit the selected cell. While editing the code, press <kbd>Enter</kbd> to create a new line, or <kbd>Shift</kbd>+<kbd>Enter</kbd> to run the code. If you are not editing the code, select a cell and press <kbd>Ctrl</kbd>+<kbd>Enter</kbd> to run the code.

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Assignment 7: Python data types: `dict`, `tuple`, `set`

In this assignment, you should write your code in a **readable** way, and **modularise** chunks of code that would otherwise be repeated.

Modularising a section of program code means to reorganise it in a way that allows it to be reused in another part of the code easily, usually in the form of a function.

Your function definitions should have appropriate **docstrings**.

## Part 1: Numerals to text

In the formal style guides for the English language, it is recommended to spell out singular numbers in full. Applications that follow this guideline have to be able to convert numbers to text in a consistent fashion.

Write a function, `text_numeral(num)` that takes in `num` (`int` < 1000) and returns a `str` of `num` in text format.

### Expected output

    >>> text_numeral(15)
    'fifteen'
    >>> text_numeral(29)
    'twenty-nine'
    >>> text_numeral(132)
    'one hundred and thirty-two'
    >>> text_numeral(250)
    'two hundred and fifty'
    
**Hint:** Draw up a list of all the words you will need. Use a `dict` to store mappings of single-word numbers to strings (e.g. one, two, tree, ... , ten, eleven, twelve, ... twenty, thirty, forty, ...).  
Numbers not in this list have to be broken down to combinations of two single-word numbers.  

<!-- grade_id: cell-ec172748d0a763bc, points: 0, grade: false, solution: true -->
```python
# Write a function to convert numbers to text numerals

def text_numeral(num:int):

```

<!-- grade_id: cell-1c54094ea29e1cf6, points: 1, grade: true, solution: true -->
## Cell for manual grading; ignore this cell

YOUR ANSWER HERE

<!-- grade_id: cell-0c7909ca3e617a23, points: 3, grade: true, solution: false -->
```python
# Run this code cell to validate your function
test_values = {15: 'fifteen',
               29: 'twenty-nine',
               132: 'one hundred and thirty-two',
               250: 'two hundred and fifty',
               735: 'seven hundred and thirty-five',
              }
for test,ans in test_values.items():
    result = text_numeral(test)
    print(f'{test}: {result}')
    assert result == ans, f'input {test} gave {result}, should be {ans} instead'
```

## Part 2: Score and grade calculation

The file `testscores.csv` contains test scores for a cohort of students. The first row contains header names.

### Task 1: Score to grade mapping

**Generate** a `dict`, `grade_for`, that has score 0-100 (`int`) as keys, and the appropriate grade as values.  
You may use a `for` or `while` loop.

The letter grade for each score range is as follows:

    A:      score >= 70
    B: 70 > score >= 60
    C: 60 > score >= 55
    D: 55 > score >= 50
    E: 50 > score >= 45
    S: 45 > score >= 40
    U: 40 > score
    
### Expected output

    >>> grade_for[30]
    'U'
    >>> grade_for[55]
    'C'
    >>> grade_for[69]
    'B'

<!-- grade_id: cell-08b4caba74070014, points: 0, grade: false, solution: true -->
```python
# Generate grade_for, a dict containing score as key and grade as value

```

<!-- grade_id: cell-206387614b7b43e3, points: 3, grade: true, solution: false -->
```python
# Run this cell to validate your dict
assert grade_for[30] == 'U', f'Grade for 30 should be U'
assert grade_for[55] == 'C', f'Grade for 55 should be C'
assert grade_for[69] == 'B', f'Grade for 69 should be B'
```

<!-- grade_id: cell-7d1f75943a2cd036, points: 2, grade: true, solution: false -->
```python
# Hidden tests; ignore this cell
```

<!-- grade_id: cell-e12a9743d19ddc85, points: 2, grade: true, solution: true -->
## Cell for manual grading; ignore this cell

YOUR ANSWER HERE

### Task 2: File reading and score calculation

Write a function, `read_testscores(filename)` that opens `filename` and:

1. Reads in the student data and returns it as a `list` of `dict`s, with each `dict` representing a student,
2. Stores the score data of each student in a `dict` with appropriate keys,
3. Calculates the overall score of each student and stores it under an `'overall'` key,
4. Determines the grade of each student and stores it under the `'grade'` key.
5. Return the `list` of student data

### Overall score calculation formula

The overall score is calculated using the following formula:

    overall = p1/30*15 + p2/40*30 + p3/80*35 + p4/30*20
    
Where `p1`, `p2`, `p3`, and `p4` are the scores for P1, P2, P3, and P4 respectively.  
The overall score is to be **rounded up** to the nearest integer. You may use the `ceil()` function from the `math` module to do this.

**Hint:** You are recommended to store the scores for P1 to P4 as a quadruple (4-`tuple`) under a `'score'` key, instead of under separate `'p1'` to `'p4'` keys, for easier retrieval.  
You may use the `zip()` function to handle multiple iterable collections in a single loop.

### Expected output

    >>> studentdata = read_testscores('testscores.csv')
    >>> studentdata[0]['class']
    'Class1'
    >>> studentdata[0]['name']
    'Student1'
    >>> studentdata[0]['overall']
    51
    >>> studentdata[0]['grade']
    'D'

<!-- grade_id: cell-4f04c23647096d3e, points: 0, grade: false, solution: true -->
```python
# Write a function that opens a file and returns student data

def read_testscores(filename):

```

<!-- grade_id: cell-3b1d7481c7ed87c9, points: 5, grade: true, solution: true -->
## Cell for manual grading; ignore this cell

YOUR ANSWER HERE

<!-- grade_id: cell-532cf5963a408736, points: 2, grade: true, solution: false -->
```python
# Run this cell to validate your dict
studentdata = read_testscores('testscores.csv')

for k,v in {'class':'Class1',
            'name': 'Student1',
            'overall': 51,
            'grade': 'D'}.items():
    assert studentdata[0][k] == v, f'Student 0: key {k} should give value {v} instead of {studentdata[0][k]}'
```

## Part 3: Grade analysis

Using `set`s, determine:

1. Which class(es) have no distinctions (grade 'A')
2. Which class(es) have 100% pass (grades A-E)

Print the result.

**Hint:** You should construct a new collection to hold the class grades, then use `set` operators to answer the above questions.

<!-- grade_id: cell-129b553507ccc54c, points: 1, grade: true, solution: true -->
```python
# Use sets to determine which class(es) have no distinctions


```

<!-- grade_id: cell-493ec62e26090c99, points: 1, grade: true, solution: true -->
```python
# Use sets to determine which class(es) have 100% pass


```

# Feedback and suggestions

Any feedback or suggestions for this assignment?

<!-- grade_id: cell-63f8f1cb7ff65bb1, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE