## 1-minute introduction to Jupyter ##

A Jupyter notebook consists of cells. Each cell contains either text or code.

A text cell will not have any text to the left of the cell. A code cell has `In [ ]:` to the left of the cell.

If the cell contains code, you can edit it. Press <kbd>Enter</kbd> to edit the selected cell. While editing the code, press <kbd>Enter</kbd> to create a new line, or <kbd>Shift</kbd>+<kbd>Enter</kbd> to run the code. If you are not editing the code, select a cell and press <kbd>Ctrl</kbd>+<kbd>Enter</kbd> to run the code.

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Assignment 6: Standard libraries, `while` loops

In this assignment, you should write your code in a **readable** way, and **modularise** chunks of code that would otherwise be repeated.

Modularising a section of program code means to reorganise it in a way that allows it to be reused in another part of the code easily, usually in the form of a function.

Your function definitions should have appropriate **docstrings**.

## Part 1

Write a function, `clock(n)`, that will show the time and update it once every second, for `n` number of seconds.

Your clock should show the time in `HH:MM:SS` format without milliseconds.

### Expected output

    >>> clock(3)
    12:59:17

<!-- grade_id: cell-7cc4baaf0efdd0cc, points: 2, grade: true, solution: true -->
```python
# Write a function to show the time and update it once every second for n seconds

### YOUR ANSWER HERE
def clock(n):


clock(3)
```

## Part 2

The `os.listdir(directory)` function from the `os` module returns a `list` of files in the given `directory`. If no argument is given to `os.listdir()`, it uses `'.'` as the default value.

    >>> import os
    >>> os.listdir('.')
    ['.ipynb_checkpoints',
     'assignment_06.ipynb',
     'folder1',
     'words.txt']
     
    >>> os.listdir()
    ['.ipynb_checkpoints',
     'assignment_06.ipynb',
     'folder1',
     'words.txt']
     
    >>> os.listdir('folder1')
    ['file1.txt', 'folder1-1', 'folder1-2']
    

### Task 1

This program code defines a function `list_files(filepath)` that takes in a `filepath` (string) and returns a `list` of files in the filepath.

    def list_files(filepath):
        import os

        list_of_files = os.listdir(filepath)
        return list_of_files

Modify the code so that it is able to return a list of files in the current directory when no argument is given to `list_files()`.

### Expected output

    >>> list_files()
    ['assignment_06.ipynb', 'folder1', '.ipynb_checkpoints']

```python
import os
os.getcwd()
```

Output:
```
ERROR:root:Internal Python error in the inspect module.
Below is the traceback from this internal error.

Traceback (most recent call last):
  File "/opt/tljh/user/lib/python3.6/site-packages/IPython/core/interactiveshell.py", line 3319, in run_code
    exec(code_obj, self.user_global_ns, self.user_ns)
  File "<ipython-input-7-a93ffb0f98e0>", line 2, in <module>
    os.getcwd()
FileNotFoundError: [Errno 2] No such file or directory

During handling of the above exception, another exception occurred:

Traceback (most recent call last):
  File "/opt/tljh/user/lib/python3.6/site-packages/IPython/core/interactiveshell.py", line 2034, in showtraceback
    stb = value._render_traceback_()
AttributeError: 'FileNotFoundError' object has no attribute '_render_traceback_'

During handling of the above exception, another exception occurred:

Traceback (most recent call last):
  File "/opt/tljh/user/lib/python3.6/site-packages/IPython/core/ultratb.py", line 1151, in get_records
    return _fixed_g
```

<!-- grade_id: cell-23b717520f88e8dc, points: 0, grade: false, solution: true -->
```python
# Task 1: Modify the code so that this code cell works without raising an Exception
### YOUR ANSWER HERE


# Do not modify any code after this line
list_files() # This line should work without modification
```

<!-- grade_id: cell-5b4e4050fa3152cd, points: 1, grade: true, solution: false -->
```python
import os
assert list_files() == os.listdir(os.getcwd()), \
'File list from list_files() does not match current directory listing'
```

### Listing subdirectories with `os.listdir()`

The path of files and folders in subdirectories must be given relative to the current folder. 

Each level down into a subfolder is demarcated by a forward slash '/' in the `filepath`.

E.g. `folder1/folder1-1` means that `folder1` is in the current directory, while `folder1-1` is a folder in `folder1`.

    >>> os.listdir('folder1-1')
    FileNotFoundError: [WinError 3] The system cannot find the path specified: 'folder1-1'
    >>> os.listdir('folder1/folder1-1')
    ['file1-1-1.txt', 'folder1-1-1']



```python
# Test your code here
os.listdir('folder1-1')
```

### Task 2

The function `walk_listing(filepath)` prints all the names of items in `filepath`. For each item, its type (file or directory) is indicated in `[`square brackets`]`.

    def walk_listing(filepath):
        list_of_files = list_files(filepath)
        for this in list_of_files:
            if os.path.isfile(this):
                print(f'[file] {this}')
            elif os.path.isdir(this):
                print(f'[dir] {this}')

Modify the function so that it **returns a list** all the names of items in `filepath`, **including the contents of subdirectories**. you do not need to include `[file]` or `[dir]`.

### Expected output

    >>> walk_listing('.')
    ['./assignment_06.ipynb',
     './folder1',
     './folder1/folder1-2',
     './folder1/folder1-2/file1-2.txt',
     './folder1/folder1-2/.ipynb_checkpoints',
     './folder1/.ipynb_checkpoints',
     './folder1/folder1-1',
     './folder1/folder1-1/file1-1-1.txt',
     './folder1/folder1-1/.ipynb_checkpoints',
     './folder1/folder1-1/folder1-1-1',
     './folder1/folder1-1/folder1-1-1/.ipynb_checkpoints',
     './folder1/folder1-1/folder1-1-1/file1-1-1-1.txt',
     './.ipynb_checkpoints',
     './.ipynb_checkpoints/assignment_06-checkpoint.ipynb']

Remember that `os.listdir()` needs the path of items in subdirectories relative to the current folder.

<!-- grade_id: cell-07e78fcbd441e896, points: 0, grade: false, solution: true -->
```python
# Task 2: Write code to list all files in the current directory recursively
import os, time

def walk_listing(filepath):
    ### YOUR ANSWER HERE
    newlist = []
    initiallist = os.listdir(filepath)
    for filename in initiallist:
        fullname = f'{filepath}/{filename}'
        newlist.append(fullname)
    i = 0
    while i < len(newlist):
        time.sleep(0.1)
        filename = newlist[i]
#     for filename in os.listdir(filepath):
        fullname = f'{filepath}/{filename}'
        newlist.append(fullname)
        if os.path.isdir(fullname):
            for subfilename in os.listdir(fullname):
                subfullname = f'{fullname}/{subfilename}'
                newlist.append(subfullname)
        i += 1
        print(newlist)

walk_listing('.')
```

<!-- grade_id: cell-eabc1493975e9724, points: 1, grade: true, solution: false -->
```python
# Run this cell to test your solution
for file in ('./assignment_06.ipynb', './folder1'):
    assert file in walk_listing('.'), f'Missing file {file} in listing'
```

<!-- grade_id: cell-982ad207e4a7b11f, points: 1, grade: true, solution: false -->
```python
# Hidden tests, ignore this cell
```

## Part 3: Lowest common multiple (LCM) with `while` loop

A simple algorithm for calculating the lowest common multiple of `a` and `b` is as follows:

1. Check if a and b are equal
2. If they are equal, a = b = lcm (lowest common multiple) [EXIT]
3. If they are not equal, check whether a or b is smaller.
4. The smaller variable is incremented by its original value.
    - i.e. a += a_original, or b += b_original
5. Go back to Step 1

Implement the above algorithm using a `while` loop. Write your code as a function, `lcm(a,b)` that returns the lowest common multiple of `a` and `b`.

### Example output

    >>> lcm(2,3)
    6
    >>> lcm(12,5)
    60

<!-- grade_id: cell-54b83892fb1327a4, points: 0, grade: false, solution: true -->
```python
def lcm(a, b):
    ### YOUR ANSWER HERE
    

lcm(2,3)
```

<!-- grade_id: cell-c5f6becdce13a059, points: 1, grade: true, solution: false -->
```python
# Run this cell to test your solution
testdata = [(2,3,6),
            (12,5,60),
            (11,12,132),
            (0,7,0),
           ]

for a,b,ans in testdata:
    assert lcm(a,b) == ans, f'Wrong result for lcm({a},{b}): should be {ans}'
```

<!-- grade_id: cell-665aea77873dd09f, points: 1, grade: true, solution: false -->
```python
# Hidden tests, ignore this cell
```

## Part 4: Greatest common factor (GCF) with `while` loop

[Euclid's Algorithm](https://en.wikipedia.org/wiki/Euclidean_algorithm) is an algorithm for calculating the greatest common factor `g` of `a` and `b` such that both `a` and `b` are divisible by `g`:

1. Check if a and b are equal
2. If they are equal, a = b = gcf (greatest common factor) [EXIT]
3. If they are not equal, the difference between the two variables, `d`, must also be divisible by `g`.
4. Calculate the GCF of the two smallest variables among `a`, `b`, and `g` (Start from Step 1)

Implement the above algorithm using a `while` loop. Write your code as a function, `gcf(a,b)` that returns the greatest common factor of `a` and `b`.

### Example output

    >>> gcf(60,48)
    12
    >>> gcf(70,14)
    14

<!-- grade_id: cell-7e2d4d1903b626da, points: 0, grade: false, solution: true -->
```python
def gcf(a, b):
    ### YOUR ANSWER HERE
    

gcf(60,48)
```

<!-- grade_id: cell-bea6eb02680a1e06, points: 1, grade: true, solution: false -->
```python
# Run this cell to test your solution
testdata = [(60,48,12),
            (70,14,14),
            (99,0,99),
#             (0,3,0),
            (0,3,3),
            (58,26,2),
            (27,13,1),
            (100,100,100),
           ]

for a,b,ans in testdata:
    result = gcf(a,b)
    assert result == ans, f'Wrong result for gcf({a},{b}): got {result}, should be {ans}'
```

<!-- grade_id: cell-9c49f2217ea939de, points: 1, grade: true, solution: false -->
```python
# Hidden tests, ignore this cell
```

# Feedback and suggestions

Any feedback or suggestions for this assignment?

<!-- grade_id: cell-63f8f1cb7ff65bb1, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE