## 1-minute introduction to Jupyter ##

A Jupyter notebook consists of cells. Each cell contains either text or code.

A text cell will not have any text to the left of the cell. A code cell has `In [ ]:` to the left of the cell.

If the cell contains code, you can edit it. Press <kbd>Enter</kbd> to edit the selected cell. While editing the code, press <kbd>Enter</kbd> to create a new line, or <kbd>Shift</kbd>+<kbd>Enter</kbd> to run the code. If you are not editing the code, select a cell and press <kbd>Ctrl</kbd>+<kbd>Enter</kbd> to run the code.

On a tablet, use the ▶️ button to run the code.

# Lesson 8b: Recursion

The `os` module in Python contains functions for working with files and directories. The `os.listdir(path)` function returns a `list` of files and folders in in the given `path`.

```python
import os

os.listdir()
```

Unfortunately, if it comes across any directories, `listdir()` will not do us the favour of listing its contents. How would we write a function that is able to get us the contents of the folder, **and** any subfolders within it?

Our function would have to:

1. Check each item in the current directory
2. Add it to the listing
3. If the item is a directory, get a list of items within the directory and add them to the listing

This would be complicated with a `for` or `while` loop. It shouldn't be this difficult! After all, our function already returns a list of items within itself and its subfolders; can't we just call this function again with every directory we encounter?

This idea of *using a function within itself* is called **recursion**.

Let's try writing this function, `recursive_listdir()`.

```python
import os

def recursive_listdir(filepath):
    listing = []
    for item in os.listdir(filepath):
        listing.append(item)
        listing.extend(recursive_listdir(item))
    return listing

recursive_listdir('.')  # recursively list the current directory
```

## Problem 1

First problem: not everything is a directory.

### Task

Modify the code below to only call `recursive_listdir()` on the item if it is a directory.

You may wish to revisit Lesson 6b if you are stuck.

```python
import os

def recursive_listdir(filepath):
    listing = []
    for item in os.listdir(filepath):
        listing.append(item)
        listing.extend(recursive_listdir(item))
    return listing

recursive_listdir('.')  # recursively list the current directory
```

Did you manage to get a listing successfully? Great!

But notice that `folder1-1` and `folder1-2` are actually not empty. Did you function include the contents of these two folders in the listing?

If yes, wonderful! If not, tackle Problem 2 below.

## Problem 2

Remember that `listdir()` only gives the names of the files or directories in the given path; it will not construct the proper path to that file/directory for you!

You may wish to revisit Lesson 6b to see how to join the path and item name to obtain the full path before passing it to `recursive_listdir()`.

## Checking a recursive implementation

In general, recursive functions should meet the following 3 conditions:

1. A recursive function should have at least one **base case** that returns a value directly without calling itself.
   - This is necessary because it must eventually reach a starting value from which it can "work backward".


2. A recursive function should call itself (**self-invocation**) when the base case is unmet.


3. Each successive recursive call should **reduce the input** (problem) to bring it *closer to the base case*.

If you are having trouble writing or formulating a recursive function, check that it fulfills all three conditions above.

# How does recursive_listdir() meet these conditions?

1. **Base case**: There must necessarily be a folder that contains no other folders.
2. **Self-invocation**: `recursive_listdir()` calls itself when a subdirectory is detected. Each subdirectory does not contain items listed in earlier recursive calls, thus shrinking the input.
3. **Return value**: `recursive_listdir()` returns the list of files and directories that it has built up so far, allowing the calling function to add that to its own list.

## (optional) Tracking recursion level

Edit the code for `recursive_listdir()` to include one more parameter, `lvl`, that indicates how many folders deep the current recursive call is in.

## Exercise 1: Write a recursive Fibonacci function

In the Fibonacci sequence, each term is the sum of the **previous two terms**.

A non-recursive implementation of the fibonacci sequence is shown below:

```python
def fibonacci(n):
    if type(n) != int or n <= 0:
        raise ValueError('n must be a positive integer (got {n})')
    if n == 1:
        return 0
    elif n == 2:
        return 1
    else:
        prev = 0
        this = 1
        for i in range(3, n + 1):
            new = this + prev
            prev = this
            this = new
        return this

for n in range(1, 10):
    print(f'{fibonacci(n)}, ', end='')
print(fibonacci(10))
```

Write a function that calculates the Fibonacci sequence **recursively**.

<!-- grade_id: cell-5bd358484bf17945, points: 0, grade: false, solution: true -->
```python
def fibonacci(n):
    if type(n) != int or n <= 0:
        raise ValueError('n must be a positive integer (got {n})')
    if n == 1:
        return 0
    elif n == 2:
        return 1
    else:
        return fibonacci(n - 1) + fibonacci(n - 2)

for n in range(1,10):
    print(f'{fibonacci(n)}, ', end='')
print(fibonacci(10))
```

<!-- grade_id: cell-d66d94a748265c3a, points: 2, grade: true, solution: false -->
```python
test_ans = [0, 1, 1, 2, 3, 5, 8, 13, 21, 34]

for i in range(len(test_ans)):
    result = fibonacci(i + 1)
    assert result == test_ans[i], \
        f'Term {i} of fibonacci sequence should be {test_ans[i]}, got {result} instead.'
```