# Lesson 5a: Functions

In lesson 4, we looked at how to work with collections of things in the form of a list. This list can be used in many ways: we can perform operations on the items inside, use it to store data from a file, and so on.

Many of the procedures we wrote in Assignments 1-4 can only be used once. If we need to use them again, we will have to copy all the code and paste it where we need it. and if we change any part of the procedure, we will have to change it in **all** the places where it is used!

A common piece of advice to Python programmers is **Don't Repeat Yourself**. Once you find yourself using the same chunk of code more than twice, you should consider how to write it in a way that you only need to declare it once and then reuse it again afterwards.

We can do so using **functions**.

## Declaring a function

You have learned how to write `if` and `for` statements:

- they start with a **keyword**
- and end with a **colon**
- all code to be executed within the block should be **indented**

Function definitions follow the same rules. A function is defined with the `def` keyword. This is how we declare a function that determines if a given `string` can be converted to a `float`:  
(Run the code cell below.)

```python
def isfloat(string):
    if ('.' in string) and (string.count('.') == 1):
        return True
    else:
        return False
    
print(isfloat('1.0'))
```

### Function input values

The function will often need to work on some input values. These input values are known as **arguments**. The arguments are passed to the function through `(` parentheses `)`. In the above function, the argument `string` is passed to the function `isfloat()` and is available for use within the function. If more than one argument needs to be passed to the function, these arguments should be separated by commas (see [Exercise 2](#Exercise-2:-Fill-in-the-blanks)).

### Function output

For some functions, an explicit return value is required. The above function needs to return a `True` value if the `string` can be converted to a `float`, and to return a `False` value if it cannot be converted to a `float`.

This is done using the `return` keyword. When the `return` keyword is invoked, Python halts execution of the function and returns the specified value to the originating statement. In the above case, the originating statement is the `print()` function.

### Exercise 1: rearrange the lines

Rearrange the following lines of code to produce a valid function declaration for the function `sigfig(num)` which takes in a single integer argument `num`.

1.  ```
        return len(sig_digits)
    ```
2.  ```
    def sigfig(num):
    ```
3.  ```
        sig_digits = num.lstrip('0').replace('.','') #Strip leftmost 0s and remove decimal
    ```
4.  ```
        num = str(num)
    ```
    
Copy and paste the lines of code into the code cell below in the appropriate order to declare a function `sigfig(num)` that returns the number of significant figures of the input number `num`. Remember to use appropriate indentation.

<!-- grade_id: cell-ad62110c33dd74fc, points: 0, grade: false, solution: true -->
```python
# Paste the code from the lines above in the correct order

### BEGIN SOLUTION
def sigfig(num):
    num = str(num)
    sig_digits = num.lstrip('0').replace('.','') #Strip leftmost 0s and remove decimal
    return len(sig_digits)
### END SOLUTION
```

<!-- grade_id: cell-01b16f46061d776c, points: 1, grade: true, solution: false -->
```python
assert sigfig(3.0) == 2, 'Incorrect order of lines'
### BEGIN HIDDEN TESTS
assert sigfig('0123') == 3, 'Incorrect order of lines'
### END HIDDEN TESTS
```

### Exercise 2: Fill in the blanks

In the code cell below, replace the underscores (`_____`) with appropriate expressions to declare a function `nth_sf(num,n)` that takes in two arguments `num` and `n` and returns the `n`th significant figure of `num`.

### Example output

```
>>> nth_sf(376.3287,1)
3
>>> nth_sf(376.3287,2)
7
>>> nth_sf(376.3287,3)
6
>>> nth_sf(376.3287,4)
3
>>> nth_sf(376.3287,5)
2
>>> nth_sf(376.3287,6)
8
```

<!-- grade_id: cell-67d4acf84d60f215, points: 1, grade: true, solution: true -->
```python
_____ nth_sf(num,n)_____
    num = str(num)
    sig_digits = num.lstrip('0').replace('.','') #Strip leftmost 0s and remove decimal
    n_sf = sig_digits[n-1] #Index starts from 0
    _____ n_sf
### BEGIN SOLUTION
def nth_sf(num,n):
    num = str(num)
    sig_digits = num.lstrip('0').replace('.','') #Strip leftmost 0s and remove decimal
    n_sf = sig_digits[n-1] #Index starts from 0
    return n_sf
### END SOLUTION
```

### Implicit `return` value

In functions where there is no `return` statement, a default value of `None` is returned.

The `None` value will be covered in later lessons.

Run the code cell below to see how a function without a `return` statement still returns a value (of `None`):

```python
def print_minutes(seconds):
    mins,sec = divmod(seconds,60)
    if mins == 1:
        s_min = ''
    else:
        s_min = 's'
    if sec == 1:
        s_sec = ''
    else:
        s_sec = 's'
        
    print(f'{mins} minute{s_min} and {sec} second{s_sec}')
    
return_value = print_minutes(61) #This assigns the return value from print_minutes() to return_value
print(return_value)
type(return_value)
```

## Don’t Repeat Yourself

A commonly quoted principle in Python programming is **Don’t Repeat Yourself**. While it is just a guideline and not to be religiously followed, often we find ourselves repeating chunks of code in multiple places in a large programming project. This makes updating the code difficult; each change we make to this chunk of code needs to be repeated in multiple places. Missing out even one occurrence can lead to a difficult-to-trace bug!

In the above code cell, there was a chunk of code that decides whether to use the plural form for units of time. This chunk of code was repeated twice to check if plural form was needed for minutes, and for seconds.

We can make the code shorter and easier to read by putting this chunk of code into another function, `suffix_for()`.

```python
def suffix_for(num):
    if num == 1:
        s = ''
    else:
        s = 's'
    return s

def print_minutes(seconds):
    mins,sec = divmod(seconds,60)
    s_min = suffix_for(mins)
    s_sec = suffix_for(sec)
        
    print(f'{mins} minute{s_min} and {sec} second{s_sec}')
    
print_minutes(62)
```

In the function `suffix_for()`, notice that the variable `s` is only used as a **temporary** holder that is immediately returned. We could simplify the code further by just returning the values `''` or `'s'` directly.

In the function `print_minutes()`, the variables `s_min` and `s_sec` function similarly as **temporary** holders that are immediately `print`ed. We could simply substitute the expressions into the f-string directly.

The following code illustrates these changes:

```python
#Improved code

def suffix_for(num):
    if num == 1:
        return ''
    else:
        return 's'

def print_minutes(seconds):
    mins,sec = divmod(seconds,60)
    print(f'{mins} minute{suffix_for(mins)} and {sec} second{suffix_for(sec)}')
    
print_minutes(121)
```

### A note on code readability

In compressing code this way, a balance has to be struck for best readability. Unnecessarily long code is tedious to read and difficult to follow, but over-compressed code can also be difficult to read.

Can you understand this code easily?

    def nth_sf(num,n):
        return str(num).lstrip('0').replace('.','')[n-1]

Code that is hard to understand at a glance can slow down an entire programming team. This is bad, especially during crunch time, when it is important to be able to quickly understand what a chunk of code is supposed to do.

Your skill at writing readable code will improve with experience. The best way to know if your code is readable is to show it to other programmers and see if they can understand it quickly and easily.

### Function declaration order

In the menu bar above, click on <kbd>Kernel</kbd> → <kbd>Restart & Clear Output</kbd> before you proceed.

Will the following code return an error? Why?

```python
print_minutes(121)

def print_minutes(seconds):
    mins,sec = divmod(seconds,60)
    print(f'{mins} minute{new_suffix_for(mins)} and {sec} second{new_suffix_for(sec)}')
    
def new_suffix_for(num):
    if num == 1:
        return ''
    else:
        return 's'
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-1-92bab4632983> in <module>
----> 1 print_minutes(121)
      2 
      3 def print_minutes(seconds):
      4     mins,sec = divmod(seconds,60)
      5     print(f'{mins} minute{new_suffix_for(mins)} and {sec} second{new_suffix_for(sec)}')

NameError: name 'print_minutes' is not defined
```

Functions must be declared before they are used. Since the Python interpreter reads the code line by line, functions are only added for use when the interpreter encounters the function declaration and **parses** it.

If you attempt to use the function `print_minutes()` before it is declared, Python checks in its memory for available functions named `print_minutes`, finds nothing, and raises a `NameError`.

### A function is an object in Python

Python treats a function as just another object. That means you can examine it with the built-in helper functions `dir()`, `help()`, and `type()`.

Try this in the code cell below and observe the output produced by the helper functions.

<!-- grade_id: cell-4beb3675e2d73f17, points: 1, grade: true, solution: true -->
```python
def print_minutes(seconds):
    '''
    Converts seconds to minutes & seconds and prints the output.
    
    Example:
    >>> print_minutes(61)
    1 minute and 1 second
    '''
    mins,sec = divmod(seconds,60)
    print(f'{mins} minute{new_suffix_for(mins)} and {sec} second{new_suffix_for(sec)}')

#Type your code below to examine the function print_minutes() with the built-in helper functions
### BEGIN SOLUTION
help(print_minutes)
### END SOLUTION
```

### Docstrings

The multiline comment (starting and ending with `'''`) you see in the previous cell, immediately below the `def` statement, is known as a **docstring**, short for **documentation string**. This docstring helps other programmers understand how to use the function. It is also returned by the `help()` function, so that programmers do not have to read the source code to know how to use it in Python.

It is standard programming practice to include docstrings in **all** functions that you write.

Docstrings may also begin and end with three double-quotes i.e. `"""`.

### Function type

Python treats functions as a special type of object. You can even write functions that take in other functions as input values!

Notice that a function has many special (dunder) methods and attributes. Almost all Python objects have special methods and attributes associated with them. We will explore special methods and attributes in future lessons.

## Variable scoping

Variables work differently inside a function and outside of the function. That is because the Python space inside the function and outside is different. We refer to these different spaces as the **scope**.

Run the code cell below and see what happens.

```python
var1 = 1
def fn():
    '''A test function to examine local vs global scoping.'''
    var1 = 2
    var2 = 2
    print('In function fn(), after declaring var2:')
    print(f'var1 is {var1}')
    print(f'var2 is {var2}')
    
fn()
```

In the above code cell, `var1` was defined first and assigned to the value `1`, outside of any functions or objects. We refer to such variables as **global** variables. Global variables exist in the global scope.

Inside the function `fn()`, the variable `var1` is assigned to the value `2`. This space is known as the **local** scope. We say that the variable `var1` has a value of `2` in the local scope and a value of `1` in the global scope.

This means that when we access `var1` in the local scope, it will return a value of `2`, and when we access `var1` in the global scope, it will return a value of `1`. You can see this happening in the code cell below, when we evaluate `var1` in the **global** scope:

```python
var1 = 1
def fn():
    '''A test function to examine local vs global scoping.'''
    var1 = 2
    var2 = 2
    print('In function fn(), after declaring var2:')
    print(f'var1 is {var1}')
    print(f'var2 is {var2}')
    
fn()
print('Outside function fn():')
print(f'var1 is {var1}')
print(f'var2 is {var2}')
```

Notice that `var2` was defined within the function, in the local scope. After the function exits and we are back in the global scope, `var2` is no longer accessible. Attempting to access `var2` in the global scope results in a `NameError` being raised.

### Protip: avoid using Python reserved keywords as variable names

This is why you should avoid using Python reserved keywords (such as `def`, `int`, `float`, `str`, ...) as variable names. You may end up overwriting important functions or keywords and not be able to use them!

In the example below, I *naively* define a variable `int` and assign it a value of `2`. This results in two `int`s existing in the global scope. Calling `int` will return the **variable** instead of the **integer** object.

```python
# This line removes any old `int` variables so that the below code cells
# do not screw up
if 'int' in globals():
    del globals()['int']

# At this point, int is a function object
print(f'int is {int}')

# Before defining an int variable, let's assign the int function to
# another object named int_fn first:
int_fn = int

# Now let's define an int variable and assign it the value 2
int = 2

# If we attempt to access the int object, we now get the int variable
# instead of the function:
print(f'int is {int}')

# fortunately, the original function is still available as int_fn:
print(f'int_fn is {int_fn}')
```

If I try to use `int()` to cast a string to an integer, that will not longer work!

```python
int('1')
```

Luckily, with foresight we had assigned the built-in `int` object to `int_fn`, so we can still use `int_fn()` to perform the casting:

```python
int_fn('1')
```

### Protip: if you absolutely must use a reserved keyword, put an underscore after it

PEP 8, the Style Guide for Python Code, [recommends using a trailing underscore_ for variables that will conflict with Python reserved keywords](https://www.python.org/dev/peps/pep-0008/#id36). For instance, to avoid a variable name collision with the reserved `class` keyword:

```
name = 'Wu Moyan'
class_ = '2021'
```

## (Optional) Nested functions

In [Don't Repeat Yourself](#Don't-Repeat-Yourself), I demonstrated how to package repeated code into a function. The functions were all defined in the **global** scope, so all the functions have access to other functions in the global scope.

What if we want to declare a function that can **only be used within another function**?

In the code cell below, `suffix_for()` is defined *within* the `print_minutes()` function. It can only be accessed in the **local** scope of the `print_minutes()` function, but not outside of it.

```python
# Remove any existing functions named print_minutes or suffix_for
if 'print_minutes' in globals():
    del globals()['print_minutes']
if 'suffix_for' in globals():
    del globals()['suffix_for']

def print_minutes(seconds):
    '''
    Prints the time given in seconds in minutes and seconds format.
    '''
    def suffix_for(num):
        '''
        Returns an appropriate suffix for num.

        Example:
        >>> suffix_for(1)
        ''
        >>> suffix_for(2)
        's'
        '''
        if num == 1:
            return ''
        else:
            return 's'
    mins,sec = divmod(seconds,60)
    print(f'{mins} minute{suffix_for(mins)} and {sec} second{suffix_for(sec)}')
    
print_minutes(121)
print(suffix_for(2)) # Cannot be accessed outside of print_minutes()
```

# Feedback and suggestions

Any feedback or suggestions for this assignment?

<!-- grade_id: cell-cff3e51d3e9a7c7b, points: 0, grade: true, solution: true -->
